# <div style="background-color:rgba(204, 229, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(0, 76, 153, 1);">PHYS 231 Lab #3<span style="color:red"> $\to$ (9 possible marks)</span></span></div>

<font size ="4"><span style="color:purple">We recommend that you complete this notebook using either the [Google Chrome](https://www.google.com/intl/en_ca/chrome/) or [Mozilla Firefox](https://www.mozilla.org/en-CA/firefox/new/) browser.  Chrome and Firefox are both available for Windows, Mac, and Linux operating systems.  The PHYS 121 Jupyter notebooks should work in other browsers, but we have specifically verfied that they work well in both Chrome and Firefox.</span></font>

# $RC$ Circuit Time Constants

***
## Lab Learning Outcomes:
By the end of the PHYS 231 Lab #3, students will be able to:
* <b><span style="color:rgba(0, 153, 76, 1);">Design, build, and test practical $RC$ circuits.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Determine the $RC$ time constant of charging and discharging capacitors using an oscilloscope and a multimeter.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Perform a weighted linear fit to a set of data and use the extracted best-fit parameters to determine a quantity of interest.</span></b>

***
## Autograding:
The PHYS 231 Labs will make use of some autograding.  To make the autograding work, the two cells below needs to be executed.  The first one installs the required packages and the second imports required packages/modules.  If 'PHYS231.Installer()' reports that some functions have been installed, the user should restart the "kernel" and then exectute the PHYS231.Installer() cell a second time.  The second time the installer function is run, it should report that **"All packages already installed. Please proceed"**.

If necessary, the kernel can be restarted by selecting **Kernel** $\to$ **Restart Kernel** from the menu options at to the top of the page.  Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/restartKernel.gif">GIF</a> showing how to restart the kernel.

The 'PHYS231.Installer()' command requires the file 'PHYS231.py', which you should see included in the list of files along the left-hand side of the screen.

```python
# Import PHYS121.py and then run the installer function.
import PHYS231
PHYS231.Installer()
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("PHYS 231 - Lab 3.ipynb")
```

***
## Import Modules:
Execute the cell below to import a number of useful pre-built Python modules that will be used in the PHYS 121 pre-labs and labs.

```python
import numpy as np
import matplotlib.pyplot as plt
from matplotlib.pyplot import cm # used to generate a sequence of colours for plotting
import pandas as pd
import scipy.stats as sp
from scipy.optimize import curve_fit
from IPython.display import HTML as html_print
from IPython.display import display, Markdown, Latex
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Name & Student Number</span></div>  

To ensure that both you and your partner receive credit for you lab work, both your name and student number and your partner's name and student number need to appear in the notebooks used during the labs.  In the cells below this text, enter your name and student number and your partner's name and student number as strings (i.e. place your names and student numbers between single or double quotes.

Examples of valid entries are:
```python
last_name_1 = 'Doe'
first_name_1 = 'Jane'
student_number_1 = '12345678'
```
and
```python
last_name_1 = "Doe"
first_name_1 = "John"
student_number_1 = "12345678"
```

<!-- BEGIN QUESTION -->

***
### **<span style="color:blue">Partner #1</span>** **<span style="color:red">(0.5 marks)</span>** 

```python
# Enter LAST name of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
last_name_1 = ...
```

```python
# Enter FIRST name of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
first_name_1 = ...
```

```python
# Enter STUDENT NUMBER of partner 1.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
student_number_1 = ...
```

***
### **<span style="color:blue">Partner #2</span>** **<span style="color:red">(0.5 marks)</span>** 
If you are working on your own, leave all of last_name_2, first_name_2, and student_number_2 as ...  Please still exectute these cells and then run the 'grader.check("name_and_student_number_2")' cell.

```python
# Enter LAST name of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
last_name_2 = ...
```

```python
# Enter FIRST name of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
first_name_2 = ...
```

```python
# Enter STUDENT NUMBER of partner 2.  Delete the ... below and type your entry between the single quotes.
# Once you've completed your entry, press SHIFT + ENTER to execute the cell.
# *** Please do not change anything to the left of the equals sign. ***
student_number_2 = ...
```

```python
grader.check("name_and_student_number")
```

<!-- END QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 1 - Measurement of the $RC$ Time Constant with the DMM (15</span></div>

Enter your measurements for the time, voltage, and uncertainty in the voltage.  For this lab, we will enter data as NumPy arrays.  Using arrays will allow us to do calculations on the dataset later on in the notebook.  The syntax is as follows:
```python
ListA = np.array([1, 2, 3, 6, 8, 12])
```
All three arrays should have the same number of elements.

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.1</span>** **<span style="color:red">(1 mark)</span>**  

In the cell below, replace the ... with your data.

*** Please do not change anything to the left of the equals sign. ***

```python
t = np.array([...])
V = np.array([...])
dV = np.array([...])
```

```python
grader.check("q2.1")
```

<!-- END QUESTION -->

### Calculate the Natural Logarithm of the Capacitor Voltage

For a capacitor discharging through a resistor, the capacitor voltage is given by:

\begin{align}
V=V_0e^{-t/\tau}
\end{align}

where $V_0$ is the initial voltage across the capacitor at time $t=0$ and $\tau = RC$ is the time constant.  This expression can be linearized by taking the natural logarithm of both sides such that:

\begin{align}
\ln\left(V\right)=\left(-\frac{1}{\tau}\right) t + \ln\left(V_0\right)
\end{align}

This is the equation of a straight line with slope $-1/\tau$ and $y$-intercept $\ln\left(V_0\right)$.

In the cell below, calculate the natural logarithm of the capacitor voltage.  Recall that in Python, the natural logarithm of $x$ is called using
```python
np.log(x)
```

***
**<span style="color:blue">Question 2.2</span>** **<span style="color:red">(1 mark)</span>**  

In the cell below, replace the ... an expression that calculates $\ln(V)$.

*** Please do not change anything to the left of the equals sign. ***

```python
lnV = ...
```

```python
grader.check("q2.2")
```

### Calculate the uncertainty in $\ln\left(V\right)$

Recall the propagation of errors.  If $y = f\left(x1, x2, ..., x_N\right)$ and $x_i\pm\Delta x_i$ for $i= 1...N$ are known, then:

\begin{align}
\Delta y = \sqrt{\left(\frac{\partial f}{\partial x_1}\Delta x_1\right)^2 + \left(\frac{\partial f}{\partial x_2}\Delta x_2\right)^2 + ... + \left(\frac{\partial f}{\partial x_N}\Delta x_N\right)^2}
\end{align}

If $y=\ln\left(V\right)$, then $y$ is a function of only a single variable and:

\begin{align}
\Delta y = \left\vert \frac{\partial y}{\partial V}\Delta V\right\vert
\end{align}

In the cell below, calculate the uncertainty in $\ln V$.

***
**<span style="color:blue">Question 2.3</span>** **<span style="color:red">(1 mark)</span>**  

In the cell below, replace the ... an expression that calculates the uncertainty in $\ln(V)$.

*** Please do not change anything to the left of the equals sign. ***

```python
errlnV = ...
```

```python
grader.check("q2.3")
```

### Plot the data

The $\ln\left(V\right)$ data can now be plotted as a function of time.  Execute the cell below to created a scatter plot with error bars.

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.4</span>** **<span style="color:red">(1 mark)</span>**  

In the cell below, replace the ... with the appropriate data to generate a plot of $ln\left(V\right)$ versus $t$.  Be sure to include errror bars.

For additional information about the 'Scatter()' function, see **Appendix A** near the end of this notebook. 

*** Please do not change anything to the left of the equals sign. ***

```python
PHYS231.Scatter(xData = ..., 
                yData = ..., 
                yErrors = ...,
                xlabel = 'time', 
                ylabel = 'ln V', 
                xUnits = 's', 
                yUnits = '');
```

<!-- END QUESTION -->

### Prepare the data for a linear fit

The scatter plot that you generated likely has a region that is linear and parts that deviate from the expected linear behaviour.  For our analysis, we want to fit only the linear portion of the data.

If you started the data collection slightly before disconnecting the power supply from your circuit, you may have several dataa points at the earliest times that are all near $10\rm\ V$.  These points will need to be removed before the fit.

It's likely that the data points at the latest times deviate from the linear trend.  This occurs because, at these times, the capacitor voltage is very small and the DMM doesn't have the resolution necessary to measure these small voltages accurately.  These points also need to be removed from the data set before doing the linear fit. 

Fortunately, Python has a very easy way to select a subset of the values contained in a vector.  The notation is as follows:

```python
x = [0, 1, 2, 3, 4, 5, 6, 7, 8, 9]

start = 2
stop = 7
y = x[2:7]
```

"start" and "stop" are positive integers.  Using the values given above, the $y$ vector will contain:

```python
y = [2, 3, 4, 5, 6]
```

Examine your plot above and adjust the "start" and "stop" values in the cell below until you have plot that looks linear.

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.5</span>** **<span style="color:red">(1 mark)</span>**  

In the cell below, replace the ... with the appropriate 'start' and 'stop' integers.

*** Please do not change anything to the left of the equals sign. ***

```python
# Enter suitable values for start and stop.
start = ...
stop = ...
t_cut = t[start:stop]
lnV_cut = lnV[start:stop]
errlnV_cut = errlnV[start:stop]

PHYS231.Scatter(
    xData = t_cut,
    yData = lnV_cut,
    yErrors = errlnV_cut,
    xlabel = "time",
    ylabel = "ln V",
    xUnits = "s",
    yUnits = "",
);
```

<!-- END QUESTION -->

## Linear Fits

The final thing that we'll do in Part 2 is fit a straight line through the data corrected (i.e. cut) data.  We'll use the 'LinearFit' contained in PHYS231.py to complete this task.  The function call to perform a weighted fit is:
```python
m, b, dm, db, fig = PHYS231.LinearFit(v_t, v_y, v_erry)
```
A weighted fit with a formatted plot can be called by using some of the other available arguments:
```python
m, b, dm, db, fig = PHYS231.LinearFit(xData = v_t, yData = v_y, yErrors = v_erry, xlabel = 'time', ylabel = 'position', xUnits = 's', yUnits = 'm')
```
Execute the cell below to preform a weighted linear fit to your data.

For additional information about the 'LinearFit' function, see **Appendix B** near the end of this notebook. 

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.6</span>** **<span style="color:red">(1 mark)</span>**  

In the cell below, replace the ... with the data needed to a weighted linear fit to $\ln\left(V\right)$ versus $t$.  Fit only the data points to follow the expected linear trend.

*** Please do not change anything to the left of the equals sign. ***

```python
m, b, dm, db, fig3 = PHYS231.LinearFit(
                    xData = ...,
                    yData= ...,
                    yErrors = ...,
                    xlabel = "time",
                    ylabel = "ln V",
                    xUnits = "s",
                    yUnits = "",
)
```

<!-- END QUESTION -->

## Determining the time constant: $\tau\pm\Delta\tau$

Use the slope from the linear fit and its uncertainty to determine an experimental value for the time constant and its uncertainty.

***
**<span style="color:blue">Question 2.7</span>** **<span style="color:red">(2 marks)</span>**  

In the cell below, replace the ... with expressions that calculte the time constant $\tau$ and its uncertainty.

*** Please do not change anything to the left of the equals sign. ***

```python
# Enter expressions to calculate tau and its uncertainty. 
tau = ...
errtau = ...

# Print tau and its uncertainty formatted to four decimal places.
print('𝜏 =', '{:.4f}'.format(tau), '+/-', '{:.4f}'.format(errtau), 's')
```

```python
grader.check("q2.7")
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 3 - Feedback and Submission</span></div>

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 3.1:</span>**  

We welcome your feedback on the PHYS 231 labs!  Please feel free to include any comments you have about this lab in the cell below.  Your comments will be taken into consideration when revising/improving the PHYS 231 labs.  You can suggest improvements, point out anything that was unclear, comment on the strengths and weaknesses of the lab, ...

This question is optional and will have no impact on your lab grade.

***
**<span style="color:blue">Answer 3.1:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell to enter your comments.  This text can be deleted.

<!-- END QUESTION -->

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

***

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Once you've completed this notebook:
- Save your work.
- Run 'grader.export()' to generate a .zip file containing all of the materials that you will submit.
- Download the generated .zip file.
- Upload the .zip file to the PHYS 121 Canvas gradebook.

Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/Submission.gif">GIF</a> showing how these steps are completed.  Once your completed notebook has been uploaded to the Canvas gradebook, you're done!

```python
# Save your notebook first, then run this cell (place your cursor in the cell and then hit Shift + Enter) to export 
# your submission.  Uploaded the .zip file that is generated to the gradebook of the Canvas shell for the PHYS 121 lab.
grader.export()
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 4 - Playground (optional)</span></div>

Feel free to add as many cells as you like below and use them as a playground for further independent investigations.  These cells won't be graded, so feel free to use them in any way that you like.  

```python
# Here's an empty code cell that you can use.
```

```python
# Here's another empty code cell that you can use.
```

```python
# Here's yet another empty code cell that you can use.
# If you need more, you can add cells using the '+' icon in the menu bar at to the top of the screen.
```

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix A &ndash; The Scatter Function...</div>

The function for generating scatter plots is called as follows:
```python
Scatter(xData, yData, yErrors = [], xlabel = 'x-axis', ylabel = 'y-axis', xUnits = '', yUnits = '', fill = False, show = True)
```
The 'xData' and 'yData' inputs are required, all other arguments are optional with default values set.  The function returns the a single output (the formatted plot):
```python
fig
```

The function will do a simple scatter plot if no 'yError' are included.  It will included error bars if 'yErrors' are passed to the function.  The 'fill' and 'show' arguments should generally be false.  

If the 'xData' list is empty, as in:
```python
xData = []
```
then the x-axis will be the trial number.  For example, is 'xData' is empty and 'yData' has 10 elements, the x-axis will span 1 to 10.

### Scatter Example Implmentation
The code block below shows an implementation of 'Scatter'.
```python
import PHYS231
theta = [10, 20, 30] # degrees
T = [2.02, 1.95, 2.13] # s
errT = [0.02, 0.02, 0.03] # s
fig = PHYS231.Scatter(theta, T, errT, 'initial angle' , 'period', 'degrees', 's')
```

If you're interested in generating your own plots with customized formatting, see the following Python-based plotting tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Basic%20Plots.html.
***

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix B &ndash; The LinearFit Function...</div>

The linear-fit function is called as follows:
```python
LinearFit(xData, yData, yErrors = [], xlabel = 'x-axis', ylabel = 'y-axis', xUnits = '', yUnits = '', fill = False)
```
The xData and yData lists are required.  The others are optional with default values set.  The function returns the following outputs:
```python
Slope, Yintercept, errSlope, errYintercept, fig
```
The first four outputs are numerical values and 'fig' is the formatted plot. 

If $y$-uncertainities are provided, the function will perform a weighted fit.  The $y$-uncertainties list must be the same length as $x$- and $y$-data lists.  If $y$-uncertainties are not provided, the fit will be unweighted.  The other optional arguments include $x$- and $y$-axis names and units for the $x$- and $y$-datasets.  These must be entered as strings (enclosed in quotations) and they are used for formatting the outputs of the function.  If 'fill' is set to True, shading will be added around the best-fit line representing confindence intervals. 

### Unweighted Fit Example Implmentation
The code block below shows the most basic use of 'LinearFit' for an unweighted fit without any of the optional arguments.
```python
import PHYS121
V = [1, 2, 3, 4]
I = [0.12, 0.198, 0.285, 0.412]
m, b, dm, db, fig = PHYS121.LinearFit(V, I)
```

### Weighted Fit Example Implmentation
This second blcok of code shows how to use 'LinearFit' to do a weighted fit.  It also makes use of all of the other optional arguments.  
```python
import PHYS121
V = [1, 2, 3, 4]
I = [0.12, 0.198, 0.285, 0.412]
errI = [0.005, 0.012, 0.020, 0.025]
m, b, dm, db, fig = PHYS121.LinearFit(V, I, errI, 'voltage' , 'current', 'V', 'A', True)
```

The 'm, b, dm, db, fig = ...' syntax is used to store the function outputs in varibles.  If you're interested in details of the fitting process, see the following Python-based tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Weighted%20Linear%20Fit.html
***

<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/images/ubc-logo-full.jpg" width="500"/>

Last update: September 28, 2023

Tests for name_and_student_number (points: None)
```
>>> isinstance(last_name_1, str)
True
```
failure_message: last_name_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(first_name_1, str)
True
```
failure_message: first_name_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(student_number_1, str)
True
```
failure_message: student_number_1 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(student_number_1, str):
...     test = len(student_number_1) == 8
>>> test
True
```
failure_message: student_number_1 must be a string of exactly 8 characters.
```
>>> test = True
>>> if isinstance(student_number_1, str):
...     test = student_number_1.isdigit()
>>> test
True
```
failure_message: student_number_1 must contain only interger digits (0, 1, 2, 3, ..., 9)
```
>>> isinstance(last_name_2, str) or last_name_2 == ...
True
```
failure_message: last_name_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(first_name_2, str) or first_name_2 == ...
True
```
failure_message: first_name_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> isinstance(student_number_2, str) or student_number_2 == ...
True
```
failure_message: student_number_2 must be a string (i.e. surround your answer by either single or double quotes).
```
>>> test = True
>>> if isinstance(student_number_2, str):
...     test = len(student_number_2) == 8
>>> test
True
```
failure_message: student_number_2 must be a string of exactly 8 characters or be left as ...
```
>>> test = True
>>> if isinstance(student_number_2, str):
...     test = student_number_2.isdigit()
>>> test
True
```
failure_message: student_number_2 must contain only interger digits (0, 1, 2, 3, ..., 9) or be left as ...


Tests for q2.1 (points: None)
```
>>> len(t) == len(V)
True
```
failure_message: Your t and V lists are different lengths.
```
>>> len(t) == len(dV)
True
```
failure_message: Your t and dV lists are different lengths.
```
>>> len(V) == len(dV)
True
```
failure_message: Your V and dV lists are different lengths.


Tests for q2.2 (points: None)
```
>>> all(np.isclose(lnV, np.log(V), rtol = 0.01))
True
```
failure_message: Double check your calculation of ln(V).  Remember that, in Python, the syntax for the natural logarithm of x is np.log(x).


Tests for q2.3 (points: None)
```
>>> all(np.isclose(errlnV, dV/V, rtol = 0.01))
True
```
failure_message: Double check your calculation of ln(V).  Remember that, in Python, the syntax for the natural logarithm of x is np.log(x).


Tests for q2.7 (points: None)
```
>>> test = True
>>> if tau > 0:
...     test = True
... else:
...     test = False
>>> test
True
```
failure_message: Check the sign of your calculated value of tau.
```
>>> np.isclose(tau, -1/m, rtol = 0.01)
True
```
failure_message: Double check the relationship between tau and the slope m.
```
>>> if errtau > 0:
...     test = True
... else:
...     test = False
>>> test
True
```
failure_message: Check the sign of your calculated value of errtau.
```
>>> np.isclose(errtau, dm/m**2, rtol = 0.01)
True
```
failure_message: Double check the relationship between errtau, the slope m, and the error in the slope dm.
